---
title: Hello World in Python
description: Have the computer talk to you for the first time, and learn about functions.


---
## Printing "Hello world!" to the console.

```yaml
type: NormalExercise
key: c493c6d860
lang: python
xp: 100
skills: 2
```


`@instructions`

Under "script.py" to the right, try typing ```print('Hello world!')``` in the console.

`@hint`

both ```'Hello world!'``` and ```"Hello world!"``` will work just the same in Python.

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}
print('Hello world!')

```

`@sct`
```{python}

```

---
## 

```yaml
type: MultipleChoiceExercise
key: 6809bf8b83
lang: python
xp: 50
skills: 2
```


`@instructions`
In Python, both ```'Hello world!'``` and ```"Hello world!"``` work the same. Knowing this, which of the following would be invalid in python?

`@hint`

`@pre_exercise_code`
```{python}

```

`@sct`
```{python}

```


